---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

You need to read in your count data from your alignemnts. This can be done in any ways. For example, you can use the xx to xxx

`@instructions`
read in your data   sfdsfdfsdfdsfdsf

`@hint`


`@pre_exercise_code`
```{r}
require(tidyverse)
```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
